```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04A.ipynb")
```

```python
from resources.hashutils import *
```

---

<h3><center>E7 -  Introduction to Programming for Scientists and Engineers</center></h3>

<h2><center>Lab session #04-A <br></center></h2>

<h1><center>Strings<br></center></h1>

---

# Question 1: Bank statement

In this problem we will create a bank statement from a list of transactions. The list is provided to us in the form of a string:

```python
flat_string='1,5,2024,ATM Withdrawal,w,50;1,10,2024,paycheck,d,2500;1,15,2024,groceries,w,93.85;1,20,2024,clothes,w,75;1,25,2024,entertainment,w,25.50'
```

The next few questions will guide you through the process of transforming the flat string into a nicely formatted bank statement. 

## Question 1.1: Split the string at semicolons (';') to obtain a list of transactions.

```python
transactions = flat_string.split(";")
transactions
```

Output:
```
['1,5,2024,ATM Withdrawal,w,50',
 '1,10,2024,paycheck,d,2500',
 '1,15,2024,groceries,w,93.85',
 '1,20,2024,clothes,w,75',
 '1,25,2024,entertainment,w,25.50']
```

```python
grader.check("q1p1")
```

Output:
```
q1p1 results: All test cases passed!
```

## Question 1.2: Split the first transaction into its parts.

Again use the string's [`split`](https://docs.python.org/3.3/library/stdtypes.html#str.split) method to split the first entry in `transactions` into 6 separate strings. Call these `month`, `day`, `year`, `itemtype`, `w_or_d`, and `amount`.

```python
month,day,year,itemtype,w_or_d,amount = transactions[0].split(",")
```

```python
grader.check("q1p2")
```

Output:
```
q1p2 results: All test cases passed!
```

## Question 1.3: Convert the day, month, and year to integers

Create new names `day_int`, `month_int`, and `year_int` for the numerical versions of these three strings.

```python
day_int = int(day)
month_int = int(month)
year_int = int(year)
```

```python
grader.check("q1p3")
```

Output:
```
q1p3 results: All test cases passed!
```

## Question 1.4: Convert w_or_d to a boolean,

Call this boolean `is_withdrawal`. Its value should be `True` if the transaction was a withdrawal and `False` if it was deposit.

```python
transactions[0]
```

Output:
```
'1,5,2024,ATM Withdrawal,w,50'
```

```python
    if w_or_d == "w": 
        is_withdrawal = True
    
    if w_or_d == "d":
        is_withdrawal = False
```

```python
grader.check("q1p4")
```

Output:
```
q1p4 results: All test cases passed!
```

## Question 1.5: Convert the amount to a float.

Name it `amount_float`.

```python
amount
```

Output:
```
'50'
```

```python
amount_float = float(amount)
```

```python
grader.check("q1p5")
```

Output:
```
q1p5 results: All test cases passed!
```

## Question 1.6: Process all transactions. Track the balance.

Write a function that applies steps 1.2 through 1.5 to each of the transactions in the list. This function takes two inputs: the raw transactions string (`flat_string` in part 1.1), and an initial balance. The output of this function is a list of tuples. Each tuple gathers information for a particular transaction. The tuple consists of:
```python
(day_int,month_int,year_int,is_withdrawal,amount_float,balance)
```

```python
def build_transactions_list(flat_string,initial_balance):

    # initialize the balance
    balance = initial_balance

    # split the input string
    transactions = flat_string.split(";")

    # create an empty list
    transactions_list = []

    for transaction in transactions:

        # split the entry
        month,day,year,itemtype,w_or_d,amount = transaction.split(",")

        # type conversions
        day_int = int(day)
        month_int = int(month)
        year_int = int(year)
        for w_or_d in transaction:
            if w_or_d == "w":
                is_withdrawal = True
            
            if w_or_d == "d":
                is_withdrawal = False
        
        amount_float = float(amount)
        
        # update the balance
        if is_withdrawal == True:
            balance -= amount_float
        if is_withdrawal == False:
            balance += amount_float

        # append this information to transactions_list
        
        transaction_tuple = day_int,month_int,year_int,is_withdrawal,amount_float,balance
        transactions_list.append(transaction_tuple)

    return transactions_list
```

```python
grader.check("q1p6")
```

Output:
```
q1p6 results: All test cases passed!
```

## Question 1.7: Format the bank statement. 

The final task is to create a nicely formatted statement to present to the customer. Your job is to write a function that transforms a list of transactions built with the `build_transactions_list` function into a list of strings as shown below. 

<img src="resources/bank_statement.png" width="800" />

The statment has these lines:

1. The name line has the string "Name:" followed by a space and the name of the customer. 
2. An empty string.
3. A separator line consisting of 53 dashes `-`.
4. A header. The parts of the header are:
    + The word `Date`, right justified in a string of width 12.
    + The word `Deposit`, right justified in a string of width 12.
    + The word `Withdrawal`, right justified in a string of width 12.
    + The word `Balance`, right justified in a string of width 12.
    + Each of these are separated by a `|` character.
5. Another separator line.
6. The list of transactions. Notice the formatting of each column. Notice also that the formatting is different for deposits and withdrawals. 
7. A final separator line.

You can check your solution against the image using the `pprint` function in the cell following the next one. 

```python
def format_statement(customer_name,transactions_list):

    WIDTH = 53 # This is the total width of the statement. 
               # and the number of dashes in the separator line.

    # Initialize and empty list of lines
    lines =[]

    # line 1: customer name
    lines.append(f"Name: {customer_name}")

    # line 2: empty
    lines.append("")

    # line 3: separator (Hint: multiply a string by an integer)
    lines.append("-"* WIDTH)

    # line 4: header
    lines.append(f"|{'Date':>12}|{'Deposit':>12}|{'Withdrawal':>12}|{'Balance':>12}|")

    # line 5: another separator
    lines.append("-"* WIDTH)

    # transaction lines
    for transaction in transactions_list:

        # unpack the transaction
        day,month,year,is_withdrawal,amount,balance=transaction

        # format the date string
        datestr =f"{month}/{day:>2}/{year}"

        # create a transaction line depending on withdrawal or deposit
        if is_withdrawal == True:
            transaction_line =f"|{datestr:>12}|{'':12}|{amount:>12.2f}|{balance:>12.2f}|"
        else:
            transaction_line =f"|{datestr:>12}|{amount:>12.2f}|{'':12}|{balance:>12.2f}|"
        # append it to lines
        lines.append(transaction_line)
    # final separator
    lines.append("-"* WIDTH)

    return lines
```

```python
from pprint import pprint

customer_name = 'Ismat Vandale'
transactions_list = build_transactions_list(flat_string,1000)
lines = format_statement(customer_name,transactions_list)
pprint(lines)
```

Output:
```
['Name: Ismat Vandale',
 '',
 '-----------------------------------------------------',
 '|        Date|     Deposit|  Withdrawal|     Balance|',
 '-----------------------------------------------------',
 '|   1/ 5/2024|            |       50.00|      950.00|',
 '|   1/10/2024|     2500.00|            |     3450.00|',
 '|   1/15/2024|            |       93.85|     3356.15|',
 '|   1/20/2024|            |       75.00|     3281.15|',
 '|   1/25/2024|            |       25.50|     3255.65|',
 '-----------------------------------------------------']
```

```python
grader.check("q1p7")
```

Output:
```
q1p7 results: All test cases passed!
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1p1 results: All test cases passed!

q1p2 results: All test cases passed!

q1p3 results: All test cases passed!

q1p4 results: All test cases passed!

q1p5 results: All test cases passed!

q1p6 results: All test cases passed!

q1p7 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1p1 (points: None)
```
>>> isinstance(transactions, list) and len(transactions) == 5
True
```
```
>>> get_hash(transactions[0]) == '4011a5fbceb3001e334b4f3ea142e590' and get_hash(transactions[1]) == 'c88ca52d57ff7ef3f932c16fb9d3ce9f'
True
```


Tests for q1p2 (points: None)
```
>>> isinstance(month, str) and isinstance(day, str) and isinstance(year, str) and isinstance(itemtype, str) and isinstance(w_or_d, str) and isinstance(amount, str)
True
```
```
>>> get_hash(month) == 'c4ca4238a0b923820dcc509a6f75849b' and get_hash(day) == 'e4da3b7fbbce2345d7772b0674a318d5' and (get_hash(year) == '07811dc6c422334ce36a09ff5cd6fe71') and (get_hash(itemtype) == '49defd0f3fdc923db244fb898145e311') and (get_hash(w_or_d) == 'f1290186a5d0b1ceab27f4e77c0c5d68') and (get_hash(amount) == 'c0c7c76d30bd3dcaefc96f40275bdc0a')
True
```


Tests for q1p3 (points: None)
```
>>> isinstance(day_int, int) and isinstance(month_int, int) and isinstance(year_int, int)
True
```
```
>>> get_hash(day_int, 1) == '336669dbe720233ed5577ddf81b653d3' and get_hash(month_int, 1) == 'e4c2e8edac362acab7123654b9e73432' and (get_hash(year_int, 4) == '48325466001978e098c8d0cb216d7cef')
True
```


Tests for q1p4 (points: None)
```
>>> isinstance(is_withdrawal, bool)
True
```
```
>>> get_hash(is_withdrawal) == 'e4c2e8edac362acab7123654b9e73432'
True
```


Tests for q1p5 (points: None)
```
>>> isinstance(amount_float, float) and np.isclose(amount_float, 50.0)
True
```


Tests for q1p6 (points: None)
```
>>> A = build_transactions_list(flat_string, 1000)
>>> np.all(np.isclose(A[0], (5, 1, 2024, True, 50.0, 950.0))) and np.all(np.isclose(A[1], (10, 1, 2024, False, 2500, 3450.0))) and np.all(np.isclose(A[2], (15, 1, 2024, True, 93.85, 3356.15))) and np.all(np.isclose(A[3], (20, 1, 2024, True, 75, 3281.15))) and np.all(np.isclose(A[4], (25, 1, 2024, True, 25.5, 3255.65)))
True
```
```
>>> flat_string = '2,9,2022,ATM Withdrawal,w,50;2,13,2024,paycheck,d,540.45;3,19,2024,groceries,w,54.82;5,03,2024,clothes,w,75.45;7,25,2024,entertainment,w,125.50'
>>> get_hash(signif(build_transactions_list(flat_string, 1000), 6)) == '3b641790d00bbbbd9198532255073756'
True
```


Tests for q1p7 (points: None)
```
>>> lines = format_statement('Ismat Vandale', build_transactions_list(flat_string, 1000))
>>> len(lines) == 11
True
```
```
>>> lines = format_statement('Ismat Vandale', build_transactions_list(flat_string, 1000))
>>> [len(line) for line in lines] == [19, 0, 53, 53, 53, 53, 53, 53, 53, 53, 53]
True
```
```
>>> lines = format_statement('Ismat Vandale', build_transactions_list(flat_string, 1000))
>>> lines[0] == 'Name: Ismat Vandale'
True
```
```
>>> stmt = format_statement('ONWRBOIWER', build_transactions_list('2,9,2022,ATM Withdrawal,w,50;2,13,2024,paycheck,d,540.45;3,19,2024,groceries,w,54.82;5,03,2024,clothes,w,75.45;7,25,2024,entertainment,w,125.50', 3243))
>>> get_hash(stmt[0]) == '31967d1cf0bfa31a78960075d65ea84d' and len(stmt[1]) == 0 and (get_hash(stmt[2]) == get_hash(stmt[4]) == get_hash(stmt[-1]) == 'd856d631339828ebc2bde0f037fdaafe') and (get_hash(stmt[-2]) == '26cc06bd62dacaa34f0a37297f9fcc37')
True
```
```
>>> stmt = format_statement('ONWRBOIWER', build_transactions_list('2,9,2022,ATM Withdrawal,w,50;2,13,2024,paycheck,d,540.45;3,19,2024,groceries,w,54.82;5,03,2024,clothes,w,75.45;7,25,2024,entertainment,w,125.50', 3243))
>>> get_hash(stmt[0]) == '31967d1cf0bfa31a78960075d65ea84d' and get_hash(stmt[-2]) == '26cc06bd62dacaa34f0a37297f9fcc37'
True
```
